# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def find_longest_name_index(names):
    """
    Find the index of the longest name in a list of names.

    :param names: A list of names to search through.
    :type names: list of str
    :return: The index of the longest name in the list.
    :rtype: int

    Examples:
    >>> find_longest_name_index(['Alice', 'Bob', 'Charlie', 'David'])
    2
    >>> find_longest_name_index(['Ella', 'Fiona', 'Grace'])
    1
    >>> find_longest_name_index(['Henry', 'Isabella', 'Jacob', 'Liam'])
    1
    """
    name_lengths = [len(name) for name in names]
    return name_lengths.index(max(name_lengths))

assert find_longest_name_index("happy birthday to me".split(" ")) == 1
assert find_longest_name_index("enjoy class".split(" ")) == 0
assert find_longest_name_index(['when', 'what', 'where', 'how', 'who']) == 2

import doctest
doctest.run_docstring_examples(find_longest_name_index, globals(), verbose=True)


### END SOLUTION
```

Output:
```
Finding tests in NoName
Trying:
    find_longest_name_index(['Alice', 'Bob', 'Charlie', 'David'])
Expecting:
    2
ok
Trying:
    find_longest_name_index(['Ella', 'Fiona', 'Grace'])
Expecting:
    1
ok
Trying:
    find_longest_name_index(['Henry', 'Isabella', 'Jacob', 'Liam'])
Expecting:
    1
ok
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert find_longest_name_index ("happy birthday to me".split(" ")) == 1
assert find_longest_name_index ("enjoy class".split(" ")) == 0
assert find_longest_name_index (['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(find_longest_name_index, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    find_longest_name_index(['Alice', 'Bob', 'Charlie', 'David'])
Expecting:
    2
ok
Trying:
    find_longest_name_index(['Ella', 'Fiona', 'Grace'])
Expecting:
    1
ok
Trying:
    find_longest_name_index(['Henry', 'Isabella', 'Jacob', 'Liam'])
Expecting:
    1
ok
```

### 25.2 Farenheit to Celsius

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION
def fahrenheit_to_celsius(fahrenheit):
    """
    Convert temperature from Fahrenheit to Celsius.

    :param fahrenheit: Temperature in Fahrenheit.
    :type fahrenheit: float
    :return: Temperature in Celsius.
    :rtype: float

    Examples:
    >>> round(fahrenheit_to_celsius(212), 2)
    100.0

    >>> round(fahrenheit_to_celsius(32), 2)
    0.0

    >>> round(fahrenheit_to_celsius(98.6), 2)
    37.0
    """
    celsius = (fahrenheit - 32) * 5/9
    return round(celsius, 2)



if __name__ == "__main__":
    temperature_in_fahrenheit = float(input("Enter temperature in Fahrenheit: "))
    temperature_in_celsius = fahrenheit_to_celsius(temperature_in_fahrenheit)
    print(f"{temperature_in_fahrenheit}째F is equal to {temperature_in_celsius}째C.")


### END SOLUTION
```

Output:
```
Enter temperature in Fahrenheit: 100
100.0째F is equal to 37.78째C.
```

```python
assert fahrenheit_to_celsius(212) == 100.0
assert fahrenheit_to_celsius(32) == 0.0
assert fahrenheit_to_celsius(98.6) == 37.0
```

```python
import doctest
doctest.run_docstring_examples(fahrenheit_to_celsius, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    round(fahrenheit_to_celsius(212), 2)
Expecting:
    100.0
ok
Trying:
    round(fahrenheit_to_celsius(32), 2)
Expecting:
    0.0
ok
Trying:
    round(fahrenheit_to_celsius(98.6), 2)
Expecting:
    37.0
ok
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates. Your function should should take three parameters: a start date, an end date, and a diagnosis.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION
from datetime import date

def length_of_stay(admit_date=date(2019, 7, 9), discharge_date=date(2019, 7, 13), diagnosis='Observation'):
    """
    Calculate the length of stay in days for a patient.

    :param admit_date: Date of admission.
    :type admit_date: datetime.date
    :param discharge_date: Date of discharge.
    :type discharge_date: datetime.date
    :param diagnosis: Diagnosis of the patient.
    :type diagnosis: str
    :return: Length of stay in days.
    :rtype: int

    Examples:
    >>> length_of_stay(date(2019, 1, 3), date(2019, 1, 8), "COPD")
    5

    >>> length_of_stay(date(2020, 1, 5), date(2020, 1, 8), "Observation")
    0
    """
    if diagnosis != 'Observation':
        length_of_stay = (discharge_date - admit_date).days
    else:
        length_of_stay = 0
    return length_of_stay



import doctest
doctest.run_docstring_examples(length_of_stay, globals(), verbose=True)


### END SOLUTION
```

Output:
```
Finding tests in NoName
Trying:
    length_of_stay(date(2019, 1, 3), date(2019, 1, 8), "COPD")
Expecting:
    5
ok
Trying:
    length_of_stay(date(2020, 1, 5), date(2020, 1, 8), "Observation")
Expecting:
    0
ok
```

```python
assert length_of_stay(date(2019,1,3), date(2019,1,8), "COPD") == 5
assert length_of_stay(date(2020,1,5), date(2020,1,8), "Observation") == 0
```

```python
import doctest
doctest.run_docstring_examples(length_of_stay, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    length_of_stay(date(2019, 1, 3), date(2019, 1, 8), "COPD")
Expecting:
    5
ok
Trying:
    length_of_stay(date(2020, 1, 5), date(2020, 1, 8), "Observation")
Expecting:
    0
ok
```

```python
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for encid, visit in encounters.items():
    print(length_of_stay(visit['admit'],visit['discharge'],visit['diagnosis']))

# Your output should be
# 5
# 4
# 1
# 0
# 0
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python

### BEGIN SOLUTION
from datetime import date

def calculate_average_length_of_stay(encounters):
    """
    Calculate the average length of stay for a list of patient encounters.

    :param encounters: Dictionary of patient encounters with admission and discharge dates.
    :type encounters: dict
    :return: Average length of stay in days.
    :rtype: float

    Examples:
    >>> round(calculate_average_length_of_stay(encounters), 2)
    3.33
    """
    length_of_stays = []
    for encounter_id, details in encounters.items():
        if details['diagnosis'] != 'Observation':
            length_of_stay = (details['discharge'] - details['admit']).days
            if length_of_stay > 0:
                length_of_stays.append(length_of_stay)

    total_length_of_stay = sum(length_of_stays)
    average_length_of_stay = total_length_of_stay / len(length_of_stays) if length_of_stays else 0

    return average_length_of_stay


# Test the function
assert round(calculate_average_length_of_stay(encounters), 2) == 3.33



### END SOLUTION
```

```python
from datetime import date
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(round(calculate_average_length_of_stay(encounters),2) == 3.33)
```

```python
import doctest
doctest.run_docstring_examples(calculate_average_length_of_stay, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    round(calculate_average_length_of_stay(encounters), 2)
Expecting:
    3.33
ok
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION

def celsius_to_fahrenheit(celsius):
    """
    Convert temperature from Celsius to Fahrenheit.

    :param celsius: Temperature in Celsius.
    :type celsius: float
    :return: Temperature in Fahrenheit.
    :rtype: float

    Examples:
    >>> celsius_to_fahrenheit(37)
    98.6

    >>> celsius_to_fahrenheit(0)
    32.0

    >>> celsius_to_fahrenheit(100)
    212.0
    """
    fahrenheit = (celsius * 9/5) + 32
    return fahrenheit

# Test cases
assert celsius_to_fahrenheit(37) == 98.6
assert celsius_to_fahrenheit(0) == 32.0
assert celsius_to_fahrenheit(100) == 212.0


### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(celsius_to_fahrenheit, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    celsius_to_fahrenheit(37)
Expecting:
    98.6
ok
Trying:
    celsius_to_fahrenheit(0)
Expecting:
    32.0
ok
Trying:
    celsius_to_fahrenheit(100)
Expecting:
    212.0
ok
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.