# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
import json

def count_entries(json_file):
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)  # Load the JSON content into a Python dictionary
            entry_list = data.get("entry", [])  # Get the "entry" list or an empty list if not found
            return len(entry_list)  # Return the count of items in the "entry" list
    except (FileNotFoundError, json.JSONDecodeError):
        return 0  # Handle exceptions by returning 0

ALLERGIES_FILE = "allergies.json"

# Testing the function
count = count_entries(ALLERGIES_FILE)
print(count)  # This should print the number of items in the "entry" list or 0 if "entry" is not found

# Validation assertions
assert isinstance(count, int)
assert count == 4



```

Output:
```
4
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
# Put your solution here
import json

def patient_count(json_file):
    with open(json_file, 'r') as file:
        data = json.load(file)

        unique_patients = set()

        for entry in data.get("entry", []):
            resource = entry.get("resource", {})
            patient_display = resource.get("patient", {}).get("display", None)

            if patient_display is not None:
                unique_patients.add(patient_display)

        return len(unique_patients)  # Return the count of distinct patient IDs

ALLERGIES_FILE = "allergies.json"

# Testing the function
distinct_patient_count = patient_count(ALLERGIES_FILE)
print(distinct_patient_count)

# Validation assertions
assert isinstance(distinct_patient_count, int)
assert distinct_patient_count == 2


```

Output:
```
2
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
# Put your solution here
import json
from collections import defaultdict

def allergy_per_patient(json_file):
    with open(json_file, 'r') as file:
        data = json.load(file)

        patient_allergy_counts = defaultdict(int)

        for entry in data.get("entry", []):
            resource = entry.get("resource", {})
            patient_display = resource.get("patient", {}).get("display", None)

            if patient_display is not None:
                patient_allergy_counts[patient_display] += 1

        # Convert the defaultdict to a regular dictionary
        return dict(patient_allergy_counts)

ALLERGIES_FILE = "allergies.json"

# Testing the function
patient_allergy_counts = allergy_per_patient(ALLERGIES_FILE)
print(patient_allergy_counts)

# Validation assertions
assert isinstance(patient_allergy_counts, dict)
assert patient_allergy_counts == {'Paul Boal': 1, 'Jason Argonaut': 3}


```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
import json

def allergy_list(json_file):
    with open(json_file, 'r') as file:
        data = json.load(file)

        results = []

        for entry in data.get("entry", []):
            resource = entry.get("resource", {})
            patient = resource.get("patient", {})
            substance = resource.get("substance", {})
            reaction = resource.get("reaction", [{}])[0].get("manifestation", [{}])[0]

            patient_display = patient.get("display", None)
            substance_text = substance.get("text", None)
            reaction_text = reaction.get("text", None)

            if all((patient_display, substance_text, reaction_text)):
                results.append([patient_display, substance_text, reaction_text])

        return results

ALLERGIES_FILE = "allergies.json"

# Testing the function
print(allergy_list(ALLERGIES_FILE))

# Given assertions for further validation
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]


```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'], ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'], ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'], ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
import json

def allergy_reaction(json_file, patient, substance):
    with open(json_file, 'r') as file:
        data = json.load(file)

        for entry in data.get("entry", []):
            resource = entry.get("resource", {})
            patient_display = resource.get("patient", {}).get("display", None)
            substance_text = resource.get("substance", {}).get("text", None)

            reaction_text = None
            reaction = resource.get("reaction", [])

            if reaction:
                manifestation = reaction[0].get("manifestation", [{}])
                if manifestation:
                    reaction_text = manifestation[0].get("text", None)

            if all((patient_display, substance_text, reaction_text)):
                if patient_display == patient and substance_text == substance:
                    return reaction_text  # Return the reaction

    return None  # Return None if no match is found

ALLERGIES_FILE = "allergies.json"

# Testing the function
print(allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G'))

# Given assertions for further validation
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') is None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'



```

Output:
```
Hives
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

